---

# Notebook Four 

#### Python skills
 - [Lists](https://automatetheboringstuff.com/chapter4/)
 
#### Instructions (important!)

- To make grading more transparent and consistent, each function below lists how many points it is worth.
- Your notebook **must** run from start to finish without errors. If your notebook hits errors, your assignment may not be graded. You can find instructions on how to make sure your notebook runs without errors [here](https://docs.google.com/document/d/1Siu7kG0X_bG5bkr9OkN6mpnL17cXL91StPgR4vGRxgY/edit?usp=sharing)
- For this notebook, you will need to delete all of the lines that say `NotImplementedError`
- You should turn in your .ipynb file to Canvas. Please don't turn in another format like HTML or PDF. You can read instructions on how to upload an ipynb [here](https://docs.google.com/document/d/1oG4HSsaSEX9jdW1iBjunR8JmfA-sgWvA-AkPYt-1wSk/edit?usp=sharing)
- Please enter your code where it says "#YOUR CODE HERE"
- Please then delete the line that says "#YOUR CODE HERE" (once you add your code).
- You will need to [upload a csv file](https://docs.google.com/document/d/1f1DQI_QUQId4x8fFQQ59cdenAZWWq9eK6HMyYchdpQM/edit?usp=sharing) to collab to complete this exercise.

### Frontier Airlines business logic

Frontier airlines boarding groups [document](https://www.flyfrontier.com/travel/travel-info/airport-info/?mobile=true) lists four groups (see "Here's who boards our planes first"). Based on these four groups, you can divide passengers into two groups. A `boards_first` group and a`regular_boarding` group, where those in the `boards_first` board first.

In this exercise, you will run a boarding simulation. [Simutation](https://en.wikipedia.org/wiki/Simulation) is a powerful comptuational technique for learning about the real world by representing it in a computer. In this homework, you will simulate a boarding process by reading in a list of passengers arriving (via a file) and make a boarding list. Your list will need to reflect the Frontier Airlines boarding order.

You should assume that every passenger in the `boards_first` group boards before every passenger in the `regular_boarding` group. You should assume that **within each group**, passengers board in the order they first arrive.

###### Deliverables:
- Your goal is to read in a list of arriving customers and put them in a list in correct order.

#### Read in from a file 

- Start off by reading in the file `passengers.jsonl` to simulate passenger arrivals
- Code for this is provided for you
- If you are curious, read more about json lines (jsonl) [here](https://jsonlines.org/)

```python
import csv
import json
import datetime
from typing import List

# notice the type hints in the function signature below (str and List[dict])
# notice also that the input format is a .jsonl file. A jsonl file is just a list of dictionaries, one per line

def simulate_passenger_arrival(filename: str ="passengers.jsonl") -> List[dict]:
    '''
    Read in a list of passengers arriving 
    '''
    output = []
    with open(filename, "r") as inf:
        for line in inf: # loop over each line in the file
            passenger = json.loads(line)
            output.append(passenger)

    # return the output list
    return output 
```

```python
# This should look familiar to HW 3 but there are subtle differences (hint: jsonl)
# It would be a good idea to compare the file read code from this HW to HW 3
passengers = simulate_passenger_arrival()
```

```python
passengers
```

Output:
```
[{'name': 'Martha Tejada', 'boarding': 'Special Services'},
 {'name': 'Valerie Mccallum', 'boarding': 'Regular boarding'},
 {'name': 'Antonio Thomas', 'boarding': 'Zone 1 Boarding'},
 {'name': 'James Bouy', 'boarding': 'Regular boarding'},
 {'name': 'Frank Gardiner', 'boarding': 'Courtesy Boarding'},
 {'name': 'William Millhouse', 'boarding': 'Regular boarding'},
 {'name': 'Lula Taveras', 'boarding': 'Regular boarding'},
 {'name': 'Sammie Lecompte', 'boarding': 'Regular boarding'},
 {'name': 'Jeffrey Russell', 'boarding': 'Regular boarding'},
 {'name': 'John Hess', 'boarding': 'Board First'}]
```

### Examine the passenger list

- You can assume the passenger list represents the passengers in the order they arrive at the airport
- So the first passenger in the list arrives first, the second passenger arrives second ...
- How each passenger represented?

```python
passengers[0]
```

Output:
```
{'name': 'Martha Tejada', 'boarding': 'Special Services'}
```

Start off by filling out the `check_board_first` function. The function should indicate if a passenger board the plane first. Unlike previous HWs, its up to you to figure out how to fill out the function from the data and business logic. Use the type hints `: dict` and `-> bool` to help you. And use an if statment (or the `==` operator).

<!-- grade_id: cell-837176e011952c2e, points: 0, grade: false, solution: true -->
```python
# Please do not delete this cell
# fill out the `boards_first` function below

def check_board_first(passenger: dict) -> bool:
    '''This function is worth 7 points'''
    if passenger['boarding'] != "Regular boarding":
      return True
    else:
      return False
            
        #if passenger["boarding"] != "Regular boarding":  this code is to print the cust 
          #print(passenger)
for passenger in passengers:
  check_board_first(passenger)
```

```python

```

<!-- grade_id: cell-9a3bafa3db108d4e, points: 7, grade: true, solution: false -->
```python
# Please do not delete this cell

### BEGIN HIDDEN TESTS

passenger1 = {'name': 'Raymond Betz', 'boarding': 'Regular boarding'}
assert not check_board_first(passenger1), "This passenger does not board first"

passenger2 = {'name': 'Lucille Hottinger', 'boarding': 'Courtesy Boarding'}
assert check_board_first(passenger2), "This passenger does board first"

### END HIDDEN TESTS
```

<!-- grade_id: cell-48b0e1d99355ae82, points: 0, grade: false, solution: true -->
```python

def simulate_boarding(passengers: List[dict]):
    '''
    - This function is worth 2 points
    - Simulate the boarding process for an airplane
    - You should take a list of passengers as input
    - You can assume that the input lists reflects the order in which 
      passengers arrive
    - You should return a list of passengers as output, which reflects the 
      airline boarding order
    - you should use your `check_board_first` function
    - You are welcome to implement this in any way that matches the business logic
    - However, a hint is provided for you on one way to do this
    - The hint uses list concatenation
        - Read more on that here: https://automatetheboringstuff.com/chapter4/
    '''
    
    # this line declares a variable called output
    # the type of the variable is a List of dictionaries
    output: List[dict] = []
    
    # this line declares a variable called boards_first
    # the type of the variable is a List of dictionaries
    boards_first: List[dict] = []
    
    # this line declares a variable called boards_second
    # the type of the variable is a List of dictionaries
    regular_boarding: List[dict] = []  

    for passenger in passengers:
        # YOUR CODE HERE
         if check_board_first(passenger) == True:
          boards_first.append(passenger)
          output.append(passenger)
         elif check_board_first(passenger) == False:
          regular_boarding.append(passenger)
          output.append(passenger)

    
    
    return boards_first + regular_boarding
```

<!-- grade_id: cell-0ec48e0218e75400, points: 2, grade: true, solution: false -->
```python
# do not delete this cell

### BEGIN HIDDEN TESTS

passengers = simulate_passenger_arrival()
boarding_order = simulate_boarding(passengers)

# find a point in the list where 
# you split between regular and priority boarding
# count up to regular board ... 
counter = 0
for passenger in boarding_order:
    if passenger["boarding"] != "Regular boarding":
        break
    else:
        counter += 1
    assert all(i["boarding"] != "Regular boarding" for i in boarding_order[0:counter]) 
    assert all(i["boarding"] == "Regular boarding" for i in boarding_order[counter:])
### END HIDDEN TESTS
```

### Challenge problem 

- This is a challenge problem worth 1 point
- The challenge problem is a little harder
- If you get all points on the HW, then getting the challenge problem right will mean a 10/10 (i.e. an A) as opposed to 9/10  (i.e., an A-)

<!-- grade_id: cell-38380e7b177c5910, points: 0, grade: false, solution: true -->
```python
from typing import List
def before(passengerA: dict, passengerB: dict, arriving_passengers: List[dict]) -> bool:
    '''
    Inputs: 
    - Passenger A, a passenger dictionary
    - Passenger B, a passenger dictionary
    - arriving_passengers, a list of passengers **in the order they arrive**
    
    Outputs:
    - True if passenger A boards before passenger B
    
    Hint: use the other functions in this HW
    '''
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-236162e5be7731bd, points: 1, grade: true, solution: false -->
```python
# Please do not delete this cell

### BEGIN HIDDEN TESTS

passengers = simulate_passenger_arrival()
boarding_order = simulate_boarding(passengers)

for passenger_number, passenger in enumerate(boarding_order):
    for other_passenger in boarding_order[passenger_number:]:
        if passenger["name"] != other_passenger["name"]:
            assert before(passenger, other_passenger, passengers)
            assert not before(other_passenger, passenger, passengers)
### END HIDDEN TESTS

```

Output:
```
Cell In[10], line 11
      9     for other_passenger in boarding_order[passenger_number:]:
     10         if passenger["name"] != other_passenger["name"]:
---> 11             assert before(passenger, other_passenger, passengers)
     12             assert not before(other_passenger, passenger, passengers)
     13 ### END HIDDEN TESTS

Cell In[9], line 16, in before(passengerA, passengerB, arriving_passengers)
      3 '''
      4 Inputs: 
      5 - Passenger A, a passenger dictionary
   (...)
     12 Hint: use the other functions in this HW
     13 '''
     15 # YOUR CODE HERE
---> 16 raise NotImplementedError()

NotImplementedError:
```

```python

```